```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment1.ipynb")
```

# STK353 Assignemnt 1

## Instructions: 
    
   - Complete all questions 
   - remember to restart your kernel and run all cells before submitting.
   - Submit only the `ipynb` file | No need to rename it.
   - Note the -> [str] at the end of functions. This is a type hint which tells you what output your function will be returning, [str] means it will return a list of strings; def fun(arg) -> float: means the function will return a float 


**Question 1**: Write a function that takes in an integer $n$ and `returns` $i^2$ for all non-negative integers $i < n$

__Example__:

The list of non-negative integers that are less than $n = 3$ is `[0, 1, 2]`. `Return` the square of each number on a separate line gives: `[0, 1, 4]`




```python
def square_int(n) -> [float]:
    result = [i**2 for i in range(n)]
    return result
```

```python
square_int(3)
```

Output:
```
[0, 1, 4]
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Question 2 (a):** You're provided the following students with their Test 1, Test 2 and Test 3 marks.

<span style="color:darkred">__Krishna = 56, 68, 80; Arjun = 79, 98, 56; Malika = 23, 43, 69__</span>
- Create a dictionary called marks to hold this data


```python
marks = {"Krishna": [56, 68, 80] ,"Arjun": [79, 98, 56], "Malika" : [23, 43, 69]}
```

```python
grader.check("q2_a")
```

Output:
```
q2_a results: All test cases passed!
```

**Question 2 (b):** Create a function that computes the mean of every student and returns the mean as a formatted float. i.e. 
   ['68.00', '77.00', '56.00']
   
- Hint: [Please use `"%.2f" % value` string formatting to get two decimal places in your average ie. xx.xx ](https://thepythonguru.com/python-string-formatting/#:~:text=2f%20is%20a%20placeholder%20for,by%20second%20value%20i.e%20150.87612%20.)


```python
def get_avg(marks)-> [str]:
    avg_mark = [ "%.2f" % (sum(i) / len(i)) for i in marks.values()] 
    return avg_mark
```

```python
get_avg(marks)
```

Output:
```
['68.00', '77.67', '45.00']
```

```python
grader.check("q2_b")
```

Output:
```
q2_b results: All test cases passed!
```

**Question 2 (c):** Create a decorator called `append_percentage` to append the percentage sign at the end of each mean. 

**Expected output**: *['68.00%', '77.67%', '45.00%']*


```python
def append_percentage(func) -> [str]:
    return func
   # ...

@append_percentage
def get_avg(marks):
    avg_marks = { Average_marks: "%.2f" % (sum(i) / len(i)) for Average_marks, i in marks.items()}
    return list(map("{}%".format,(avg_marks.values())))
    ...
```

```python
get_avg(marks)
```

Output:
```
['68.00%', '77.67%', '45.00%']
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q2_a results: All test cases passed!

q2_b results:
    q2_b - 1 result:
        Test case passed!

    q2_b - 2 result:
        Test case passed!

    q2_b - 3 result:
        Trying:
            get_avg({'Krishna': [56, 68, 80] ,'Arjun': [79, 98, 56], 'Malika' : [23, 43, 69]}) == ['68.00', '77.67', '45.00']
        Expecting:
            True
        **********************************************************************
        Line 1, in q2_b 2
        Failed example:
            get_avg({'Krishna': [56, 68, 80] ,'Arjun': [79, 98, 56], 'Malika' : [23, 43, 69]}) == ['68.00', '77.67', '45.00']
        Expected:
            True
        Got:
            False

q2_c results: All test cases passed!
```

Tests for q1 (points: [2, 3])
```
>>> isinstance(square_int(2), list)
True
```


Tests for q2_a (points: [1, 1, 1])
```
>>> isinstance(marks, dict)
True
```


Tests for q2_b (points: 6)
```
>>> marks = {'Krishna': [1] ,'Arjun': [2], 'Malika' : [3]}
>>> assert isinstance(get_avg(marks), list)
```


Tests for q2_c (points: [2, 2])
